---
title: Chapter 1
description: Basic operations.
---

## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 2a2109b11f
```

Do some basic calculations using R

`@instructions`

Add two and two.

`@hint`

2 + 2

`@pre_exercise_code`

```{r}

```

`@sample_code`

```{r}

```

`@solution`

```{r}
# Calculate the sum of 2 and 2
2 + 2
```

`@sct`

```{r}
test_object()
```
